```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# E7: Lab Assignment 02 - Data Structures & NumPy Arrays

You must submit the lab to Gradescope by the due date. You will submit the zip file produced by running the final cell of the assignment.

## Instructions

**Run the first cell, Initialize Otter**, to import the autograder and submission exporter.

Throughout the assignment, replace `...` with your answers. We use `...` as a placeholder and these should be deleted and replaced with your answers.

Any part listed as a "<font color='red'>**Question**</font>" should be answered to receive credit.

**Please save your work after every question!**

To read the documentation on a Python function, you can type `help()` and add the function name between parentheses.

## Score Breakdown
Question   | Points
:---       | --:
1.0 – 1.9  | 6
2          | 2
3.0 – 3.5  | 3
4.0 – 4.10 | 9
5          | 2
6.0 – 6.1  | 3
Total      | 25

**Run the cell below**, to import the required modules.

```python
# Please run this cell, and do not modify the contents
import numpy as np
import math
import matplotlib.pyplot as plt
import re
import pandas as pd
np.seterr(all='ignore');
import hashlib
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Question 1: Elections

We need to analyze an election dataset which contains the names of voters and the candidates they voted for.

Run the next cell, which reads the data and saves the values as `voters` and `candidates`. In the future, we will learn how to read files in more detail. For now, use the provided command to read the file. The code cell below reads the CSV file, saves its values as a DataFrame `df`, and then extracts each column, saving them as `voters` and `candidates`.

Run the cell below and do not modify its contents.

```python
# Please run this cell, and do not modify the contents

# Read the data
df = pd.read_csv('resources/elections.csv')

# Display a sample of the data
print(df.head())

# Save each column separately
voters = df[df.columns[0]].tolist()
candidates = df[df.columns[1]].tolist()
```

### Question 1.0

What are the data types of `voters` and `candidates`? Assign them to the variables `type_voters` and `type_candidates`, respectively.

```python
# ANSWER CELL
type_voters = ...
type_candidates = ...

print(f'Type voters: {type_voters} \nType candidates: {type_candidates}')
```

```python
grader.check("q1.0")
```

### Question 1.1

What are the data types of the elements stored in `voters` and `candidates`? Assign these data types to the variables `elem_type_voters` and `elem_type_candidates`, respectively. Assume that all elements in `voters` have the same data type and all elements in `candidates` have the same data type. 

```python
# ANSWER CELL
elem_type_voters = ...
elem_type_candidates = ...

print(f'Type voters elements: {elem_type_voters} \nType candidates elements: {elem_type_candidates}')
```

```python
grader.check("q1.1")
```

### Question 1.2

What is the total number of registered voters who voted? Assign your answer to the variable `num_voters`. Do not just manually type the numeric answer. Use Python functions to determine the value.

```python
# ANSWER CELL
num_voters = ...

print(f'Number of voters: {num_voters}')
```

```python
grader.check("q1.2")
```

### Question 1.3

If there were 10,372 registered voters, what was the voter turnout? Report you answer as a percentage and assign it to variable `turnout_voters`. You should compute this value directly in Python and assign the expression that returns the answer to `turnout_voters`. Do not just manually type the numeric answer. Use Python operators to determine the value.

Note that voter turnout is defined as the percentage of registered voters who actually voted.

```python
# ANSWER CELL
turnout_voters = ...

print(f'Voter turnout {turnout_voters} %')
```

```python
grader.check("q1.3")
```

### Question 1.4

Note that the data are organized such that `voters[i]` voted for `candidates[i]` for every index `i`. Assume that the names of the voters are unique, such that the same name does not appear twice.

Who did Amy King vote for? First, use list methods to find Amy King's index in the list of voters, and then use this to find the candidate they voted for. Assign Amy King's index to `amy_king_index` and the candidate they voted for to `amy_king_vote`. Do not just manually type the answer. Use Python expressions to determine the value.

```python
# ANSWER CELL
amy_king_index = ...
amy_king_vote = ...

print(f'Candidate name: {amy_king_vote}')
```

```python
grader.check("q1.4")
```

### Question 1.5

We need to write a code that generates the following statement:

‘In the recent elections, [Voter Name] voted for [Candidate Name].’

In this statement, the words in brackets will be replaced with the name of a voter and the candidate they voted for.

To achieve this, we will use functions. A function in Python is a block of reusable code that performs a specific task. It can take input, process it, and return an output. We will learn about functions in more detail next week.

For this task, you will modify a pre-made function called `VoterInfo`, which takes the name of a voter as input and outputs the desired statement.

Assume that `name` is a variable that contains a voter's name as a string, exactly as it appears in `voters`. Your goal is to assign the string that includes the output statement for any `name` to the variable `string`. Be careful with your string formatting. Any small modification, such as an extra space or character, may result in a failed test. The expression you assign to `string` should work for any value of `name`.

Here are some examples and their expected outputs:

```
>>> VoterInfo(name = 'Amy King')
'In the recent elections, Amy King voted for Lila Martinez.'

>>> VoterInfo(name = 'Laura Harrison')
'In the recent elections, Laura Harrison voted for Zach Johnson.'
```

```python
# ANSWER CELL

def VoterInfo(name):
    
    ########### DO NOT MODIFY ABOVE THIS LINE ###########
    
    # Keep string indented at the same level as other lines of code
    string = ...
    
    ########### DO NOT MODIFY BELOW THIS LINE ###########
    return string
```

```python
# TEST YOUR FUNCTION
# Feel free to try different names to confirm that your function works for any input
# Note that the name should be in the list voters. Otherwise, you will get an error

VoterInfo(name = 'Amy King')
```

```python
grader.check("q1.5")
```

### Question 1.6

How many total votes did candidate Lila Martinez get? Save this as variable`votes_LM`. Note that your answer should be an integer. Do not just manually type the answer. Use Python expressions to determine the value.

You may find the list methods we discussed in the lecture helpful.

```python
# ANSWER CELL
votes_LM = ...

print(f'Total votes Lila Martinez: {votes_LM}')
```

```python
grader.check("q1.6")
```

### Question 1.7

Create a list named `votes` that includes the total number of votes that candidates Lila Martinez, Zach Johnson, and Tyler Davis received, in this order. Do not just manually type the answer. Use Python expressions to determine the value.

```python
# ANSWER CELL
votes = ...

print(f'Total votes: {votes}')
```

```python
grader.check("q1.7")
```

### Question 1.8

Someone claims that candidate Tyler Davis received more votes than either candidate Lila Martinez or candidate Zach Johnson.

Construct a logical expression to check if this claim is true or false and assign it to `q1_8`.

```python
# ANSWER CELL
q1_8 = ...

print(f'Tyler Davis received more votes than either Lila Martinez or Zach Johnson: {q1_8}')
```

```python
grader.check("q1.8")
```

### Question 1.9

Someone claims that candidate Tyler Davis received more votes than both candidates Lila Martinez and Zach Johnson combined.

Construct a logical expression to check if this claim is true or false and assign it to `q1_9`.

```python
# ANSWER CELL
q1_9 = ...

print(f'Tyler Davis received more votes than Lila Martinez and Zach Johnson combined: {q1_9}')
```

```python
grader.check("q1.9")
```

## Question 2: Number Echo

We need to write a code that repeats an integer `n` times. For example, if `n = 5`, the output should be

‘55555’

For this task, you will modify a pre-made function called `Repeat`, which takes an integer input and outputs the desired string.

Assume that `n` is a non-negative integer. Your goal is to assign the string that includes the repeated integer `n` for any given value of `n` to the variable `string`. Be careful with your string formatting. Any small modification, such as an extra space or character, may result in a failed test. The expression you assign to `string` should work for any value of `n`.

Here are some examples and their expected outputs:

```
>>> Repeat(n = 4)
'4444'

>>> Repeat(n = 10)
'10101010101010101010'
```

*Hint: what happens when you multiply a string by an integer?*

```python
# ANSWER CELL

def Repeat(n):
    
    ########### DO NOT MODIFY ABOVE THIS LINE ###########
    
    # Keep string indented at the same level as other lines of code
    string = ...
    
    ########### DO NOT MODIFY BELOW THIS LINE ###########
    return string
```

```python
# TEST YOUR FUNCTION
# Feel free to try different n to confirm that your function works for any input

Repeat(n = 4)
```

```python
grader.check("q2")
```

## Question 3: California Reservoirs

For this question, you will analyze water storage of three reservoirs in California from Jan-Dec 2022. First, you will clean and manipulate the data and then, you will calculate basic statistics for each reservoir.

The three reservoirs you will analyze are Lake Oroville, Shasta Lake, and Hetch Hetchy Reservoir. Shasta Lake and Lake Oroville are the largest and second-largest reservoirs in California, providing flood-control, hydropower, and water storage for the state. Shasta Lake is a key reservoir within the federal government's Central Valley Project. Similarly, Lake Oroville is a major reservoir in the state government's California State Water Project. Hetch Hetchy reservoir is within Yosemite National Park and provides water to the City of San Francisco and other counties in the San Francisco Bay Area.

We will read a text file that includes the monthly storage volumes for each reservoir, taken from January through December 2022, in units of acre-feet. We will provide you with the command to read the file, but in the future, you will learn how to read text files and other file types. In the code cell below, we are reading a text file and saving its values as `reservoir_data`. Run the cell below and do not modify its contents.

```python
# Please run this cell, and do not modify the contents
reservoir_data = np.loadtxt('resources/res_vals_int.csv',delimiter='\t')
print(reservoir_data)
```

The 12 rows correspond to the monthly storage volumes and the three columns correspond to the three reservoirs in the following order:
1. Lake Oroville
2. Shasta Lake
3. Hetch Hetchy Reservoir

### Question 3.0

Create a list `reservoir_names` that includes the following two elements: `'Lake Oroville'` and `'Shasta'` (in the same order as listed here).

```python
# ANSWER CELL
...
```

```python
# Run this cell before running grader.check()
q3_0 = reservoir_names.copy()
```

```python
grader.check("q3.0")
```

### Question 3.1

In the list `reservoir_names`, Shasta Lake is listed as just `'Shasta'`. Modify `reservoir_names` to change `'Shasta'` to `'Shasta Lake'`. Do not create a new variable `reservoir_names` – use list indexing.

```python
# ANSWER CELL
...
```

```python
# Run this cell before running grader.check()
q3_1 = reservoir_names.copy()
```

```python
grader.check("q3.1")
```

### Question 3.2
Our list `reservoir_names` only has two entries! We are missing the third entry `'Hetch Hetchy Reservoir'`. Add `'Hetch Hetchy Reservoir'` to the end of the `reservoir_names` list. Do not redefine the list. Use appropriate list methods that allow you to add an entry to an existing list.

```python
# ANSWER CELL
...
```

```python
# Run this cell before running grader.check()
q3_2 = reservoir_names.copy()
```

```python
grader.check("q3.2")
```

### Question 3.3

We want to create a range of numbers 1 through 12 to represent the rows within the `reservoir_data` array (i.e., the months Jan - Dec).

Create a range of integers from 1 to 12, inclusive, using the `range` function. Store the range of numbers as `months_range`.

```python
# ANSWER CELL
...
```

```python
grader.check("q3.3")
```

### Question 3.4
Convert `months_range` to a list. Store the values as `months_list`. The variable `months_list` should contain the integers 1 through 12, just like `months_range`.

```python
# ANSWER CELL
months_list = ...

print(months_list)
```

```python
grader.check("q3.4")
```

### Question 3.5

What is the data type of `reservoir_data`? Assign it to the variable `type_reservoir_data`.

```python
# ANSWER CELL
type_reservoir_data = ...

type_reservoir_data
```

```python
grader.check("q3.5")
```

<div class="alert alert-block alert-warning"> <b>NOTE!</b> The rest of the lab involves working with NumPy arrays. If we haven't covered this topic yet and you have no prior experience with it, please wait until after the next lecture to continue the assignment. </div>

## Question 4: California Reservoirs (Continued)

Recall that the 12 rows of `reservoir_data` correspond to the monthly storage volumes and the three columns correspond to the three reservoirs in the following order:
1. Lake Oroville
2. Shasta Lake
3. Hetch Hetchy Reservoir

### Question 4.0
First, slice `reservoir_data` to select the column of data corresponding to Lake Oroville for all months.

Store the sliced data as `OrovilleData`. The output should be a vector (not a matrix) of the monthly storage values for Lake Oroville.

```python
# ANSWER CELL
...
```

```python
grader.check("q4.0")
```

### Question 4.1
Slice `reservoir_data` to select the winter storage volumes for Hetch Hetchy Reservoir. For this question, use January, February, and March as the winter months.

Store the sliced data as `WinterHetchHetchy`.

```python
# ANSWER CELL
...
```

```python
grader.check("q4.1")
```

### Question 4.2
Slice `reservoir_data` to select the storage volumes for the even months (February, April, ..., October, and December) for Lake Oroville and Hetch Hetchy Reservoir only.

Store the sliced data as `EvenMonths`.

```python
# ANSWER CELL
...
```

```python
grader.check("q4.2")
```

### Question 4.3
For each month, how much more water was stored in Lake Oroville than in Hetch Hetchy? Report your answer in thousand acre-feet (TAF).

Store the differences in thousand acre-feet as `Diff`.

*Hint*: Your answer should be an array which includes the storage difference for each month. This can be achieved using arithmetic operations on slices of `reservoir_data`. Divide by 1000 to convert from acre-feet to TAF.

```python
# ANSWER CELL
...
```

```python
grader.check("q4.3")
```

### Question 4.4
Find the mean storage volume for Summer (July, August, and September) for Shasta Lake. California receives very little rainfall in the summer months, so reservoir storage during these months is of particular interest.

Store the mean value as `SummerMeanStorage`.

*Hint*: Slice `reservoir_data` to extract the months and reservoir of interest, then calculate the mean of the sliced array. While there are functions to compute the mean directly, built-in function such as `sum()` and `len()` can also be used.

```python
# ANSWER CELL
...
```

```python
grader.check("q4.4")
```

### Question 4.5
Check if the the mean value you computed above, `SummerMeanStorage`, exists in `reservoir_data`. Use membership operators and assign the expression to `q4_5`. The expression should return `True` or `False`, depending on whether the value of `SummerMeanStorage` exists in `reservoir_data` or not.

```python
# ANSWER CELL
q4_5 = ...

print(f'The mean value exists in reservoir_data: {q4_5}')
```

```python
grader.check("q4.5")
```

### Question 4.6
For each month, did the storage volume in **Lake Oroville** exceed `SummerMeanStorage`?

You should answer this by creating a logical expression using the array `OrovilleData`. The output should be another array that includes `True` or `False` for each month, depending on whether each month had storages higher than `SummerMeanStorage` for Lake Oroville. This should be stored in `q4_6`.

```python
# ANSWER CELL
q4_6 = ...

print(f'Volume in Lake Oroville exceeded the mean: {q4_6}')
```

```python
grader.check("q4.6")
```

### Question 4.7
For each month, did the storage volume in **Lake Oroville or Shasta Lake** exceed `SummerMeanStorage`?

You should answer this by creating a logical expression using the array `reservoir_data`. The output should be another array that includes `True` or `False` for each month, depending on whether each month had storages higher than `SummerMeanStorage` for either Lake Oroville or Shasta Lake. This should be stored in `q4_7`.

```python
# ANSWER CELL
q4_7 = ...

print(f'Volume in Lake Oroville or Shasta Lake exceeded the mean: {q4_7}')
```

```python
grader.check("q4.7")
```

### Question 4.8
For how many months did the storage volume in **Lake Oroville or Shasta Lake** exceed `SummerMeanStorage`?

Store the value as `q4_8`.

*Hint:* True is equivalent to 1 and False is equivalent to 0. So, you can take the sum of the logical array from the previous question to find the total number of months for which the logical expression is `True`.

```python
# ANSWER CELL
q4_8 = ...

print(f'Volume in Lake Oroville or Shasta Lake exceeded the mean for {q4_8} months.')
```

```python
grader.check("q4.8")
```

### Question 4.9
We would like to convert the data in `reservoir_data` from acre-feet to $m^3$. 1 acre-feet is equal to 1233.48 $m^3$.

Store the value as `reservoir_data_m3`.

```python
# ANSWER CELL
reservoir_data_m3 = ...

reservoir_data_m3
```

```python
grader.check("q4.9")
```

### Question 4.10
Lastly, we want to create a NumPy array that is the same shape of `reservoir_data_m3` but has 1 for each entry where the storage is higher than 2,000,000,000 $m^3$ and 0 otherwise.

For example, if we had the storage data shown below, the output should like as follos:

$$\text{Storage: } \begin{matrix}
  [[\text{2.1e+09}, & \text{3.1e+09}, & \text{1.1e+09}] \\
  \ [\text{1.9e+09}, & \text{1.5e+09}, & \text{2.2e+09}] \\
  \ [\text{1.3e+09}, & \text{2.1e+09}, & \ \text{1.4e+09}]]
  \end{matrix}
  \rightarrow 
  \text{Output: }
  \begin{matrix}
  [[\text{1}, & \text{1}, & \text{0}] \\
  \ [\text{0}, & \text{0}, & \text{1}] \\
  \ [\text{0}, & \text{1}, & \ \text{0}]]
  \end{matrix}
$$

Assign the array to the variable `q4_10`. Do not create the array manually. Begin by creating an array that has the shape of interest then modify its elements using indexing and logical expressions. The elements inside `q4_10` should be of type `int`.

Note that NumPy arrays have a method called `astype` (`arr.astype('int')`), which converts the floating-point numbers in the array to integers.

```python
# ANSWER CELL
...
```

```python
grader.check("q4.10")
```

## Question 5: Extract Reservoir Information

Using the same dataset of water storage, we need to write a code that takes information about reservoir storage as inputs and then outputs a statement: 

‘In [Month Name], [Reservoir Name] held [Volume] acre-feet of water.’ 

In this statement, the words in brackets will be replaced with the name of the input month, reservoir name, and storage volume, respectively.

For this task, you will modify a pre-made function called `ReservoirInfo`, which takes the name of a reservoir and the month number as inputs and outputs the desired statement.

Assume that `name` is a variable that contains a reservoir's name as a string, exactly as it appears in `reservoir_names` and that `month` is an integer representing the month number, where 1 corresponds to January, 2 corresponds to February, etc. Your goal is to assign the string that includes the output statement for any `name` and `month` to the variable `string`. Be careful with your string formatting. Any small modification, such as an extra space or character, may result in a failed test. The expression you assign to `string` should work for any value of `name` and `month`.

Here are some examples and their expected outputs:

```
>>> ReservoirInfo(name = 'Lake Oroville', month = 1)
'In January, Lake Oroville held 1642000 acre-feet of water.'

>>> ReservoirInfo(name = 'Shasta Lake', month = 5)
'In May, Shasta Lake held 1815000 acre-feet of water.'
```

Important notes:
* [Month Name] should be the full month name of the input number. You may find the [`calendar.month_name[number]`](https://docs.python.org/3/library/calendar.html#calendar.month_name) attribute helpful. To use it, you have to import the module `calendar` first.
* [Volume] should be obtained from `reservoir_data` based on the corresponding input month and reservoir name. You may find the list `index()` method helpful (returns the index of the specified element in the list). Note that [Volume] should be an **integer** and not float.

```python
# ANSWER CELL

def ReservoirInfo(name,month):

    ########### DO NOT MODIFY ABOVE THIS LINE ###########
    
    # Keep string indented at the same level as other lines of code
    string = ...
    
    ########### DO NOT MODIFY BELOW THIS LINE ###########
    return string
```

```python
# TEST YOUR FUNCTION
# Feel free to try different names and months to confirm that your function works for any input

ReservoirInfo(name = 'Lake Oroville', month = 1)
```

```python
grader.check("q5")
```

## Question 6: Plot Reservoir Storage

Finally for this assignment, you will define the input data and title string for a pre-made plotting function called `plot_storage`.

### Question 6.0

Assume that `name` is a variable that contains a reservoir's name as a string, exactly as it appears in `reservoir_names`. Your goal is to define `title_string` as a format string to have the reservoir name input, `name`, followed by "storage in 2022". For example, a graph of Lake Oroville should have the title "Lake Oroville storage in 2022". The variable `title_string` should automatically update based on the input `name`. Be careful with your string formatting. Any small modification, such as extra space or character, will result in a failed test.

Then, define `data`, which should include the storage volumes for all months for the corresponding reservoir. `data` should be obtained from `reservoir_data` based on the corresponding input reservoir name. You may find the list `index()` method helpful (returns the index of the specified element in the list).

```python
# ANSWER CELL
def plot_storage(name):
    
    ########### DO NOT MODIFY ABOVE THIS LINE ###########
    
    # Keep the lines indented at the same level as other lines of code
    title_string = ...
    data = ...
    
    ########### DO NOT MODIFY BELOW THIS LINE ###########
    x_label = 'Month'
    y_label = 'Reservoir storage (acre-feet)'
    plt.plot(months_range, data)
    plt.xlabel(x_label)
    plt.ylabel(y_label)
    plt.title(title_string)
    plt.xticks(range(1, 13))
    plt.show()
    return title_string , data
```

```python
# TEST YOUR FUNCTION
# Feel free to try different names to confirm that your function works for any input

plot_storage(name = 'Lake Oroville')
```

```python
# Run this cell before running grader.check()
fig1 = plot_storage(name = 'Lake Oroville')
fig2 = plot_storage(name = 'Shasta Lake')
fig3 = plot_storage(name = 'Hetch Hetchy Reservoir')
```

```python
grader.check("q6.0")
```

### Question 6.1

The function from Q6.0, returns (sends back) both `title_string` and `data` (see last line in the function). In the code below, we are saving the output of the function as `q6_1`. Evaluate the data type and length of `q6_1` and save them as `q6_1_type` and `q6_1_length`, respectively.

```python
# ANSWER CELL

q6_1 = plot_storage(name = 'Lake Oroville')

q6_1_type = ...
q6_1_length = ...
print(f'Type: {q6_1_type} \nLength: {q6_1_length}')
```

```python
grader.check("q6.1")
```

### You're done with this Lab!

**Important submission information:** After completing the assignment, click on the Save icon from the Tool Bar &nbsp;<i class="fa fa-save" style="font-size:16px;"></i>&nbsp;. After saving your notebook, **run the cell with** `grader.check_all()` and confirm that you pass the same tests as in the notebook. Then, **run the final cell** `grader.export()` and click the link to download the zip file. Then, go to Gradescope and submit the zip file to the corresponding assignment. 

**Once you have submitted, stay on the Gradescope page to confirm that you pass the same tests as in the notebook.**

```python
import matplotlib.image as mpimg
img = mpimg.imread('resources/animal.jpg')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congrats on finishing this lab!")
plt.show()
```

```python

```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1.0 (points: 0.25)
```
>>> assert get_hash(type_voters) == '425ed5b515f0a67c6316f0ec499f86bf'
>>> assert get_hash(type_candidates) == '425ed5b515f0a67c6316f0ec499f86bf'
```


Tests for q1.1 (points: 0.25)
```
>>> assert get_hash(elem_type_voters) == '9a86641cdf2fdb47f786dc955264738d'
>>> assert get_hash(elem_type_candidates) == '9a86641cdf2fdb47f786dc955264738d'
```


Tests for q1.2 (points: 0.5)
```
>>> assert get_hash(int(num_voters)) == '0d98b597aa732aea606bde680c3b57d8'
```


Tests for q1.3 (points: 0.5)
```
>>> assert get_hash(round(turnout_voters,10)) == '608880fae9997ba4bb218db46e20d6f4'
```


Tests for q1.4 (points: 1)
```
>>> assert get_hash(amy_king_index) == '88ccf9a98e8519b2011df33952832f2f'
>>> assert get_hash(amy_king_vote.upper()) == '25e5a1adec953dde8ae67b8fa66d3369'
```


Tests for q1.5 (points: 1)
```
>>> 
>>> # Check data structure
>>> assert get_hash(len(VoterInfo(name=voters[0]))) == 'ea5d2f1c4608232e07d3aa3d998e5135'
```
failure_message: Check spacing and data type.
```
>>> 
>>> # Check the return type
>>> assert get_hash(type(VoterInfo(name=voters[0]))) == '9a86641cdf2fdb47f786dc955264738d'
```
failure_message: Check the return type.
```
>>> 
>>> # Check wording of string output
>>> assert get_hash("RECENT" in VoterInfo(name=voters[0]).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash("ELECTIONS" in VoterInfo(name=voters[0]).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash("VOTED" in VoterInfo(name=voters[0]).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Check the wording.
```
>>> 
>>> # Check reservoir name
>>> assert get_hash("MARIA WALKER" in VoterInfo(name=voters[7]).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash("DENNIS BOYD" in VoterInfo(name=voters[451]).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash("JAMES WILLIAMS" in VoterInfo(name=voters[6888]).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Check the voter name.
```
>>> 
>>> # Check the name of the month
>>> assert get_hash("TYLER DAVIS" in VoterInfo(name=voters[6201]).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash("LILA MARTINEZ" in VoterInfo(name=voters[2361]).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash("ZACH JOHNSON" in VoterInfo(name=voters[67]).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Check the candidate's name.


Tests for q1.6 (points: 0.5)
```
>>> assert get_hash(int(votes_LM)) == '43e4e6a6f341e00671e123714de019a8'
```


Tests for q1.7 (points: 1)
```
>>> assert get_hash(type(votes)) == '425ed5b515f0a67c6316f0ec499f86bf'
```
```
>>> assert get_hash(len(votes)) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```
```
>>> assert get_hash(sum(votes)) == '0d98b597aa732aea606bde680c3b57d8'
>>> assert get_hash(votes[0]) == '43e4e6a6f341e00671e123714de019a8'
>>> assert get_hash(votes[1]) == '46a558d97954d0692411c861cf78ef79'
>>> assert get_hash(votes[2]) == '2d2ca7eedf739ef4c3800713ec482e1a'
```


Tests for q1.8 (points: 0.5)
```
>>> assert get_hash(q1_8) == 'f827cf462f62848df37c5e1e94a4da74'
```


Tests for q1.9 (points: 0.5)
```
>>> assert get_hash(q1_9) == 'f8320b26d30ab433c5a54546d21f414c'
```


Tests for q2 (points: 2)
```
>>> 
>>> # Check data structure
>>> for i in range(10):
...     assert len(Repeat(i)) == i
...     
>>> for i in range(10,100):
...     assert len(Repeat(i)) == i*2
```
failure_message: Check spacing and data type.
```
>>> 
>>> # Check the return type
>>> assert get_hash(type(Repeat(n = 5))) == '9a86641cdf2fdb47f786dc955264738d'
```
failure_message: Check the return type.
```
>>> 
>>> # Check string output
>>> assert get_hash(Repeat(1)) == 'c4ca4238a0b923820dcc509a6f75849b'
>>> assert get_hash(Repeat(8)) == '8ddcff3a80f4189ca1c9d4d902c3c909'
>>> assert get_hash(Repeat(12)) == '984719247dace834974b72df7923ee1f'
>>> assert get_hash(Repeat(22)) == '67ba0218ba471508b94b07a22919454e'
>>> assert get_hash(Repeat(75)) == 'ba50dbae347c4e5a95a241bec2c0b4da'
```
failure_message: Check the output.


Tests for q3.0 (points: 0.5)
```
>>> 
>>> assert get_hash(len(q3_0)) == 'c81e728d9d4c2f636f067f89cc14862c'
```
failure_message: You should only have two values (for now).
```
>>> 
>>> assert get_hash(q3_0[0]) == 'b947db2fb4b5855984587a0d4fabf896'
```
failure_message: Check value of Lake Oroville entry.
```
>>> 
>>> assert get_hash(q3_0[1]) == '36d4e7e1bfd840adcaede22706f227f9'
```
failure_message: Check value of Shasta Lake entry.
```
>>> 
>>> assert get_hash(type(q3_0)) == '425ed5b515f0a67c6316f0ec499f86bf'
```
failure_message: Check data type of reservoir_names.


Tests for q3.1 (points: 0.5)
```
>>> 
>>> assert get_hash(len(q3_1)) == 'c81e728d9d4c2f636f067f89cc14862c'
```
failure_message: You should still only have two values.
```
>>> 
>>> assert get_hash(q3_1[0]) == 'b947db2fb4b5855984587a0d4fabf896'
```
failure_message: Do not edit the Lake Oroville entry.
```
>>> 
>>> assert get_hash(q3_1[1]) =='dd7d39986c5d807a0248838b12acc7f3'
```
failure_message: Check value of Shasta Lake entry.


Tests for q3.2 (points: 0.5)
```
>>> 
>>> assert get_hash(len(q3_2)) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```
failure_message: Check length of list.
```
>>> 
>>> assert get_hash(q3_2[-1].upper()) == 'e5bb3ed8fc109fa9f05b0d538a30edf8'
```
failure_message: Check value Hetch Hetchy entry.


Tests for q3.3 (points: 0.5)
```
>>> assert get_hash(months_range) == '5319bc85f3a29d6cd7b6b65df958ae78'
```
```
>>> 
>>> assert get_hash(months_range[0]) == 'c4ca4238a0b923820dcc509a6f75849b'
```
failure_message: Check start value.
```
>>> 
>>> assert get_hash(months_range[-1]) == 'c20ad4d76fe97759aa27a0c99bff6710'
```
failure_message: Check end value.
```
>>> 
>>> assert get_hash(sum(months_range)) == '35f4a8d465e6e1edc05f3d8ab658c551'
```
failure_message: Check step size.


Tests for q3.4 (points: 0.5)
```
>>> 
>>> assert get_hash(type(months_list)) == '425ed5b515f0a67c6316f0ec499f86bf'
>>> assert get_hash(len(months_list)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(sum(months_list)) == '35f4a8d465e6e1edc05f3d8ab658c551'
```
failure_message: Check months_list.


Tests for q3.5 (points: 0.5)
```
>>> assert get_hash(type_reservoir_data) == '30b39aca73d0815f5ca1a06af1412947'
```


Tests for q4.0 (points: 0.5)
```
>>> 
>>> assert get_hash(len(OrovilleData)) == 'c20ad4d76fe97759aa27a0c99bff6710'
```
failure_message: Check length of selection.
```
>>> 
>>> assert get_hash(int(sum(OrovilleData))) == '97c2a51b4ddc77196fa563cd72957c1e'
```
failure_message: Check which column corresponds to Lake Oroville.
```
>>> assert get_hash(OrovilleData.astype(int)) == 'd57f3682934eebd41c73b68817c6592d'
```


Tests for q4.1 (points: 1)
```
>>> 
>>> assert get_hash(len(WinterHetchHetchy)) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```
failure_message: Check length of selection.
```
>>> 
>>> assert get_hash(int(sum(WinterHetchHetchy))) == '50037bd6578a3e4db86ac430e190f7cb'
```
failure_message: Check which column corresponds to Hetch Hetchy Reservoir.
```
>>> assert get_hash(WinterHetchHetchy.astype(int)) == '7b67ebf4e3d0d34b2c62dba86b620e1b'
```


Tests for q4.10 (points: 1)
```
>>> 
>>> assert get_hash(q4_10.shape) == '52658b9c1506282ca4f20ade14e21d99'
>>> assert get_hash(type(q4_10)) == '30b39aca73d0815f5ca1a06af1412947'
```
failure_message: Check the shape and type of the output. It should be an ndarray that has the same shape as the storage data.
```
>>> 
>>> assert get_hash(type(q4_10[0,0])) != '781f5f2052cfb496fa73d70ea1fbdaa8'
```
failure_message: Check the data type of the entries.
```
>>> 
>>> assert get_hash(sum(q4_10).astype(int)) == '91f3c39cea609dafc87061b3e6686199'
>>> assert all([i == 0 for i in q4_10[:,-1]])
>>> assert all([q4_10[i,0] == 1 for i in range(len(q4_10[:,0])) if i < 6])
>>> assert all([q4_10[i,0] == 0 for i in range(len(q4_10[:,0])) if i > 5])
>>> assert all([q4_10[i,1] == 1 for i in range(1, len(q4_10[:,1])) if i < 7])
>>> assert all([q4_10[i,1] == 0 for i in range(1, len(q4_10[:,1])) if i > 6])
```
failure_message: Check the entries of the output.


Tests for q4.2 (points: 1)
```
>>> 
>>> assert get_hash(len(EvenMonths)) == '1679091c5a880faf6fb5e6087eb1b2dc'
```
failure_message: Check length of selection.
```
>>> 
>>> assert get_hash(EvenMonths.shape) == 'e186b2fc415e1b6085ca473fcc974c66'
```
failure_message: Check slice selection.
```
>>> assert get_hash(EvenMonths.astype(int)) == 'b2f43e102ce4fe6ac1033f199ea9cd4b'
```


Tests for q4.3 (points: 1)
```
>>> assert get_hash(len(Diff)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(Diff.shape) == '3f30626b8bdbdc8366257c053839456f'
>>> assert get_hash(sum(Diff)) == '4a3dd518aa065cd57d5d5acfeb1ff6da'
>>> assert get_hash(Diff) == 'bde6ec575d6f776bfcb4df36f0e64c28'
```


Tests for q4.4 (points: 1)
```
>>> 
>>> assert get_hash(round(SummerMeanStorage)) == '118df32ac7db6de47c75d4ac84e6af2e'
```
failure_message: Check mean calculation.


Tests for q4.5 (points: 0.5)
```
>>> assert get_hash(q4_5) == 'f8320b26d30ab433c5a54546d21f414c'
```


Tests for q4.6 (points: 1)
```
>>> assert get_hash(len(q4_6)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(q4_6.shape) == '3f30626b8bdbdc8366257c053839456f'
>>> assert get_hash(q4_6) == '453afeb97740796e13cd9c391a6c9c86'
```


Tests for q4.7 (points: 1)
```
>>> assert get_hash(len(q4_7)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(q4_7.shape) == '3f30626b8bdbdc8366257c053839456f'
>>> assert get_hash(q4_7) == 'bc1f55c0a039c7a0df8b2ece9ceda908'
```


Tests for q4.8 (points: 0.5)
```
>>> assert get_hash(q4_8) == '8f14e45fceea167a5a36dedd4bea2543'
```


Tests for q4.9 (points: 0.5)
```
>>> assert get_hash(len(reservoir_data_m3)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(reservoir_data_m3.shape) == '52658b9c1506282ca4f20ade14e21d99'
>>> assert get_hash(sum(reservoir_data_m3)) == '4170ddab9b71a0810a7a6b5bb1e4444c'
```


Tests for q5 (points: 2)
```
>>> 
>>> # Check data structure
>>> import calendar
>>> assert get_hash(len(ReservoirInfo(name=reservoir_names[2],month=2))) == '735b90b4568125ed6c3f678819b6e058'
```
failure_message: Check spacing and data type.
```
>>> 
>>> # Check the return type
>>> assert get_hash(type(ReservoirInfo(name=reservoir_names[1],month=5).upper())) == '9a86641cdf2fdb47f786dc955264738d'
```
failure_message: Check the return type.
```
>>> 
>>> # Check wording of string output
>>> assert get_hash("HELD" in ReservoirInfo(name=reservoir_names[1],month=5).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash("WATER" in ReservoirInfo(name=reservoir_names[1],month=5).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Check the wording.
```
>>> 
>>> # Check reservoir name
>>> assert get_hash("SHASTA" in ReservoirInfo(name=reservoir_names[1],month=8).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash("OROVILLE" in ReservoirInfo(name=reservoir_names[0],month=8).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Check the lake name.
```
>>> 
>>> # Check the name of the month
>>> assert get_hash("MAY" in ReservoirInfo(name=reservoir_names[1],month=5).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash("JULY" in ReservoirInfo(name=reservoir_names[1],month=7).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash("FEBRUARY" in ReservoirInfo(name=reservoir_names[1],month=2).upper()) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Check the month name.
```
>>> 
>>> # Check value
>>> assert get_hash("1621000" in ReservoirInfo(name=reservoir_names[1],month=1)) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash("1690000" in ReservoirInfo(name=reservoir_names[1],month=2)) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash("305100" in ReservoirInfo(name=reservoir_names[2],month=8)) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Check storage value.


Tests for q6.0 (points: 2)
```
>>> 
>>> assert get_hash(fig1[0].upper()) == '00871030319cbec6a516cff2df01cab7'
>>> assert get_hash(fig2[0].upper()) == 'ab05a7c024ce9e92c0c987edea417f6a'
>>> assert get_hash(fig3[0].upper()) == 'ad274fcd8f530d1e1b1ddfa131eded78'
```
failure_message: Check title.
```
>>> 
>>> assert get_hash(len(fig1[1])) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(len(fig2[1])) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(len(fig3[1])) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(int(sum(fig1[1]))) == get_hash(int(sum(reservoir_data[:,0])))
>>> assert get_hash(int(sum(fig2[1]))) == get_hash(int(sum(reservoir_data[:,1])))
>>> assert get_hash(int(sum(fig3[1]))) == get_hash(int(sum(reservoir_data[:,2])))
```
failure_message: Check data.


Tests for q6.1 (points: 1)
```
>>> assert get_hash(q6_1_type) == '22b3afce62075c7012f8e5041adfee16'
>>> assert get_hash(q6_1_length) == 'c81e728d9d4c2f636f067f89cc14862c'
```
