```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

# E7: Lab Assignment 04 - Iteration & Recursion

You must submit the lab to Gradescope by the due date. You will submit the zip file produced by running the final cell of the assignment.

## Instructions

**Run the first cell, Initialize Otter**, to import the autograder and submission exporter.

Throughout the assignment, replace `...` with your answers. We use `...` as a placeholder and these should be deleted and replaced with your answers.

Any part listed as a "<font color='red'>**Question**</font>" should be answered to receive credit.

**Please save your work after every question!**

To read the documentation on a Python function, you can type `help()` and add the function name between parentheses.

## Score Breakdown
Question   | Points
:---       | --:
1          | 4
2.0 – 2.2  | 8
3          | 3
4          | 4
5          | 3
6          | 3
Total      | 25

**Run the cell below**, to import the required modules.

```python
# Please run this cell, and do not modify the contents
import numpy as np
import math
import random
import hashlib
import matplotlib.pyplot as plt
import inspect
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Question 1: Interest

The interest rate, $r$, on a principal investment, $P_0$, is a payment for allowing the bank to use an investor's money. Compound interest is accumulated according to the formula $P_n = (1 + r)P_{n-1}$, where $P_n$ is the balance after $n$ compounding periods (in our case years), and $P_{n-1}$ is the balance after the previous period. The interest rate $r$ is given
in decimal, e.g. 0.1, which is 10%. The objective is to determine the number of years required to reach an investment goal based on an initial investment $P_0$ and an interest rate $r$.

You will need to complete the function below, `mySavingsPlan(P0, r, goal)`, which returns `years`, the number of years required to reach an investment goal. Assume that interest is paid out in one payment at the end of a given year. The output `years` should be a whole number.

Complete the function below to solve the investment problem using a while loop. There are hints to guide you on what is missing on each line.

Test your `mySavingsPlan()` function for `P0 = 1000` dollars, `r = 0.05`, and `goal = 10000` dollars. Assign the output to `q1`. Feel free to try it for other inputs.

```python
# ANSWER CELL
def mySavingsPlan(P0, r, goal):
    """Return number of years required to reach an investment goal based on an initial investment P0 and an interest rate r.
    >>> mySavingsPlan(1000, 0.05, 10000)
    48
    >>> mySavingsPlan(1000, 0.10, 10000)
    25
    >>> mySavingsPlan(1000, 0.05, 1000000)
    142
    """
    # Initialize number of years required to reach investment goal
    years = ...
    
    # Initialize amount of money in the current period
    money = ...
    
    # Add logical expression that should be checked each time you iterate through the loop
    while ...
        
        # Update amount of money in the next period
        money = ...
        
        # Update number of years required to reach investment goal
        years = ...

    return years
```

```python
# TEST YOUR FUNCTION HERE
q1 = ...

print(f'Number of years to reach goal: {q1} years')
```

```python
grader.check("q1")
```

## Question 2: Sum Digits

### Question 2.0

Write a function `Sum(n)` that takes a non-negative integer (`n`) and returns the sum of its digits (`digit_sum`) as an integer. If `n` is negative and/or non-integer, the function should return "Invalid Input." Be careful with your string formatting. Minor differences will result in a failed test.

Test your `Sum(n)` function for  `n = 200500` in the following cell. Assign the output to `q2_0`. Feel free to try it for other inputs.

*Hint:* Strings are iterable.

```
Examples:

>>> Sum(10) # 1 + 0 = 1
1

>>> Sum(4224) # 4 + 2 + 2 + 4 = 12
12

>>> Sum(1234567890)
45

>>> Sum(-10)
'Invalid Input.'

>>> Sum(10.1)
'Invalid Input.'

>>> a = Sum(123) # make sure that you are using return rather than print
>>> a
6
```

```python
# ANSWER CELL
...
```

```python
# TEST YOUR FUNCTION HERE
q2_0 = ...

print(f'Sum of  digits: {q2_0}')
```

```python
grader.check("q2.0")
```

### Question 2.1

Create a simpler version of Q2.0 by writing a lambda function called `lambdaSum`. Your lambda function should perform the same calculation as Q2.0 with `n` as the input and sum of the digits as the output. Also, the function should return "Invalid Input." if `n` is negative and/or non-integer.

Test your new `lambdaSum` function for `n = 200500` and assign the output to `q2_1`. Feel free to try it for other inputs.

*Hint: You might need to use list comprehension and ternary operators.*

```
Examples:

>>> lambdaSum(10) # 1 + 0 = 1
1

>>> lambdaSum(4224) # 4 + 2 + 2 + 4 = 12
12

>>> lambdaSum(1234567890)
45

>>> lambdaSum(-10)
'Invalid Input.'

>>> lambdaSum(10.1)
'Invalid Input.'
```

```python
# ANSWER CELL
...
```

```python
# TEST YOUR FUNCTION HERE
q2_1 = ...

print(f'Sum of digits: {q2_1}')
```

```python
grader.check("q2.1")
```

### Question 2.2

In the previous question, we considered one digit at a time when computing the sum. Next, we will write a similar function that computes the sum of groups of digits of a specified size.

Write a function `SumSize(n, size)` that takes a non-negative integer `n` and another integer `size` which specifies the number of digits to consider at a time. The function should return the sum of `size` consecutive digits from `n`.

For example, if `n = 200500` and `size = 2`, the function should return 20 + 05 + 00 = 25 (two consecutive digits at a time and taking their sum).

Set the default value of `size` equal to 1 and assume `n` is a valid input (you do not have to check if it is negative or float).

Note that the number of digits in `n` should be a multiple of `size`. If this is not the case, your function should return the following string:
* "The number of digits in n should be a multiple of size."

Test your `SumSize(n, size)` function for  `n = 200500` and `size = 2`. Assign the output to `q2_2`. Feel free to try it for other inputs.


```
Examples:

>>> SumSize(10) # 1 + 0 = 1 (size not specified, so use default value of 1)
1

>>> SumSize(10, 2) 
10

>>> SumSize(4224, 2) # 42 + 24 = 66
66

>>> SumSize(100001111101234, 5) # 10000 + 11111 + 01234 = 22345
22345

>>> SumSize(4224, 3)
"The number of digits in n should be a multiple of size."
```

```python
# ANSWER CELL
...
```

```python
# TEST YOUR FUNCTION HERE
q2_2 = ...

print(f'Sum of group of digits: {q2_2}')
```

```python
grader.check("q2.2")
```

## Question 3: Collatz Conjecture

The Collatz conjecture is one of the most famous unsolved problems in mathematics:

> Pick a positive integer `n`, any positive integer. If the `n` is even, divide it by 2; if it's odd, multiply it by 3 and add 1. Take that new number and repeat the process, again and again. If you keep repeating this, you'll eventually reach the number 1, regardless of which positive integer `n` is chosen initially. And if you continue the process beyond 1, you will get stuck in a loop. 

$$10 \rightarrow 5 \rightarrow 16 \rightarrow 8 \rightarrow 4 → 2 \rightarrow 1 \rightarrow 4 \rightarrow 2 \rightarrow 1 \rightarrow 4 \rightarrow 2 \rightarrow 1 \rightarrow \dots$$

$$11 \rightarrow 34 \rightarrow 17 \rightarrow 52 \rightarrow 26 \rightarrow 13 \rightarrow 40 \rightarrow 20 \rightarrow 10 \rightarrow 5 \rightarrow 16 \rightarrow 8 \rightarrow 4 \rightarrow 2 \rightarrow 1 \rightarrow 4 \rightarrow \dots $$

The number `n` will travel up and down but eventually end at 1 (at least for all numbers that have ever been tried -- nobody has ever proved that the sequence will terminate). 

Write a function `collatz(n)` that takes a positive integer (`n`) and returns two outputs:
1. `sequence`: a list of the sequence starting at `n` and ending at the first occurrence of 1
2. `steps`: the number of steps in the sequence

If `n` is non-positive and/or non-integer, the function should return "Invalid Input." Be careful with your string formatting. Minor differences will result in a failed test. 

Also note that if `n == 1` initially, then the sequence is one step long.

Test your `collatz(n)` function for  `n = 25`. Assign the output to `q3`. Feel free to try it for other inputs.

```
Examples:

>>> collatz(10)
([10, 5, 16, 8, 4, 2, 1], 7)

>>> collatz(1)
([1], 1)

>>> collatz(9)
([9, 28, 14, 7, 22, 11, 34, 17, 52, 26, 13, 40, 20, 10, 5, 16, 8, 4, 2, 1], 20)

>>> collatz(0)
'Invalid Input.'
```

```python
# ANSWER CELL
...
```

```python
# TEST YOUR FUNCTION HERE
q3 = ...

print(f'Sequence: {q3[0]}')
print(f'Number of steps to get to 1: {q3[1]}')
```

```python
grader.check("q3")
```

## Question 4: Rainfall Totals

Write a function `Rainfall(rainData)` that takes a numeric sequence (list, tuple, or NumPy array) of daily rainfall measurements, `rainData`, and returns the total rainfall, `rainTot`, and the average daily rainfall, `rainAvg`, over a period of interest.

The argument `rainData` may contain the number -999 indicating the end of the data from a given period. We are interested in computing the total and average daily rainfall over the first period only (everything after the first -999 should be ignored). There may be negative numbers other than -999 in`rainData`, which represent erroneous measurements and should be neglected.

Your function `myRainfall(rainData)` should return `rainTot` and `rainAvg`, which are the total and average of the non-negative values in either the entire input `rainData` or the period from the beginning until the first occurrence of -999. Note that `rainAvg` is the average daily rainfall during this period (not counting days with neglected measurements). The first returned object should be `rainTot` and the second returned object should be `rainAvg`.

Test your `myRainfall()` function for `rainData = [2.3, 4.1, -5.4, 1.0, 7.1, -2, 8, 0, 4.9, 6, -999]`. Assign the output to `q4`. Feel free to try it for other inputs.

```
Examples:

>>> myRainfall([2, 4, -5, 1, 7, -2, 8, -999, 4, 6, 4]) # rainTot = 2 + 4 + 1 + 7 + 8 = 22
(22, 4.4)

>>> myRainfall([2, 4, -5, 1, 7, -2, 8, 0, 4, 6, 4]) # rainTot = 2 + 4 + 1 + 7 + 8 + 0 + 4 + 6 + 4 = 36
(36, 4.0)

>>> myRainfall([2, 4, -999, 1, 7, -2, 8, -999, 4, 6, 4]) # rainTot = 2 + 4 = 6
(6, 3.0)

>>> myRainfall([0, -999, -5, 1, 7, -2, 8, -999, 4, 6, 2]) # rainTot = 0
(0, 0.0)
```

```python
# ANSWER CELL
...
```

```python
# TEST YOUR FUNCTION HERE
q4 = ...

print(f'Rainfall total: {q4[0]}')
print(f'Rainfall average: {q4[1]}')
```

```python
grader.check("q4")
```

<div class="alert alert-block alert-warning"> <b>NOTE!</b> The rest of the lab focuses on recursion. If we haven't covered this topic yet and you have no prior experience with it, please wait until after the next lecture to continue the assignment. </div>

## Question 5: Recursive Sequence

Consider the following function

\begin{equation}
F_n = \begin{cases}
    F_0 &\text{if $n=0$}\\
    F_1 &\text{if $n=1$}\\
    aF_{n-1} + bF_{n-2} & \text{if $n\geq2$}\\
    \end{cases}
\end{equation}

where $F_0, F_1, a, \text{ and } b$ are known constants.

Using **recursion**, write a function `GenRec(a, b, F0, F1, n)`, that returns the $n^{th}$ number in the sequence, $F_n$. Assume that `n` is a non-negative integer (no need to check it in your function).

Test your `GenRec()` function for any set of inputs. Assign the output to `q5`.

```
Examples:

>>> GenRec(1, 1, 0, 1, 0)
0

>>> GenRec(1, 1, 0, 1, 1)
1

>>> GenRec(1, 1, 0, 1, 6)
8

>>> GenRec(1, 2, 0, 1, 6)
21

>>> GenRec(1, 2, 0, 1, 6)
5858

>>> GenRec(2, 2, 0, 1, 6)
56320
```

```python
# ANSWER CELL
...
```

```python
# TEST YOUR FUNCTION HERE
q5 = ...

print(f'F(n) = {q5}')
```

```python
grader.check("q5")
```

## Question 6: Fibonacci but on Steroids

A *k-bonacci sequence* is a generalization of the Fibonacci sequence. In a k-bonacci sequence, each term is the sum of the k preceding terms. The Fibonacci sequence is a special case with k = 2, where each term is the sum of the two previous terms. In a 3-bonacci sequence (k = 3), each term is the sum of the three previous terms, and so on. A k-bonacci sequence starts with k $-$ 1 zeros and then a one. Each subsequent element is the sum of the previous k elements. For example:

```
        n: 0, 1, 2, 3, 4, 5, 6,  7,  8,  9, ..., n

2-bonacci: 0, 1, 1, 2, 3, 5, 8, 13, 21, 34, ..., F(n-1) + F(n-2)
3-bonacci: 0, 0, 1, 1, 2, 4, 7, 13, 24, 44, ..., F(n-1) + F(n-2) + F(n-3)
4-bonacci: 0, 0, 0, 1, 1, 2, 4,  8, 15, 29, ..., F(n-1) + F(n-2) + F(n-3) + F(n-4)
```

Hint: Using this information, try to come up with the k-bonacci formula for different values of $n$. For any k, what should be F(0)? What about F(k-1)? F(k)? F(n)? Your formula should look something like this:

\begin{equation}
F(n) = \begin{cases}
    \dots &\text{if } \dots\\
    \dots &\text{if } \dots\\
    \dots & \text{if } \dots\\
    \end{cases}
\end{equation}

Write a function `kbonacci()` which takes a non-negative integer `n` and a positive integer `k`. The function should return the value of the $n^{th}$ element of the k-bonacci sequence.

You will implement the function `kbonacci()` by filling in the blanks. You may only write one statement on each line. You may not use ; to place two statements on the same line. You may not use any built-in functions or methods. You may not modify the level of indentation of any line.

```python
def kbonacci(n, k):
    if n < k - 1:
        return ...(a)...
    elif ...(b)...:
        return 1
    else:
        total = 0
        i = ...(c)...
        while i < n:
            total += ...(d)...
            i += 1
    return total
```

<br> 

```
Examples:

>>> kbonacci(8, 3)
24 
    
>>> kbonacci(9, 3)
44

>>> kbonacci(9, 4)
29
```

### Question 6.0

Fill in blank (a). Assign the expression that should replace blank (a) to the variable `q6a`.

```python
# ANSWER CELL
q6a = ...

q6a
```

```python
grader.check("q6.0")
```

### Question 6.1

Which of these could fill in blank (b)? Assign your answer to the variable `q6b` as a string.

**A.** `n == k` \
**B.** `n = k` \
**C.** `n > k` \
**D.** `n > k - 1` \
**E.** `n < k - 2` \
**F.** `n == k - 1` \
**G.** `n = k - 1` \
**H.** `n == 0` \
**I.** `n == 1`

Your answer should be a string, e.g., `"A"`, `"B"`, etc.\
Remember to put quotes around your answer choice.

**Note that the test for this question will be a hidden test. Meaning, you will NOT be able to know whether your answer is correct or not by running the `grader.check()` cell.** However, you can add all your answers to Question 6.4 and confirm that your function gives the correct answer.

```python
# ANSWER CELL
q6b = ...

q6b
```

```python
grader.check("q6.1")
```

### Question 6.2

Which of these could fill in blank (c)? Assign your answer to the variable `q6c` as a string.

**A.** `k` \
**B.** `0` \
**C.** `i` \
**D.** `n - k` \
**E.** `k - 1` \
**F.** `1` \
**G.** `k + 1` \
**H.** `k - n` \
**I.** `k` \
**J.** `3` 

Your answer should be a string, e.g., `"A"`, `"B"`, etc.\
Remember to put quotes around your answer choice.

**Note that the test for this question will be a hidden test. Meaning, you will NOT be able to know whether your answer is correct or not by running the `grader.check()` cell.** However, you can add all your answers to Question 6.4 and confirm that your function gives the correct answer.

```python
# ANSWER CELL
q6c = ...

q6c
```

```python
grader.check("q6.2")
```

### Question 6.3

Which of these could fill in blank (d)? Assign your answer to the variable `q6d` as a string.

**A.** `1` \
**B.** `kbonacci(n-1, k)` \
**C.** `kbonacci(i, k)` \
**D.** `total + 1` \
**E.** `total + F(n-1)` \
**F.** `n` \
**G.** `i` \
**H.** `k` \
**I.** `F(i)` \
**J.** `k - 1` 

Your answer should be a string, e.g., `"A"`, `"B"`, etc.\
Remember to put quotes around your answer choice.

**Note that the test for this question will be a hidden test. Meaning, you will NOT be able to know whether your answer is correct or not by running the `grader.check()` cell.** However, you can add all your answers to Question 6.4 and confirm that your function gives the correct answer.

```python
# ANSWER CELL
q6d = ...

q6d
```

```python
grader.check("q6.3")
```

### Question 6.4

Add your answers from Questions 6.0 – 6.3 below and test your function.

```python
# ANSWER CELL
def kbonacci(n, k):
    if n < k - 1:
        return ...
    elif ...:
        return 1
    else:
        total = 0
        i = ...
        while i < n:
            total += ...
            i += 1
    return total
```

```python
# TEST YOUR FUNCTION HERE
...
```

```python
grader.check("q6.4")
```

##  (Optional) Question 7: Vwl Rmvr!

**Note: This problem is NOT graded and is provided solely for additional practice. Skipping or completing this problem will not affect your grade.**

Write a function `remove_vowels(input_string)` which takes in a single string as its argument and performs two tasks:
* Removes all uppercase and lowercase vowel characters ('A', 'E', 'I', 'O', 'U') from the input string. 
* If an exclamation sign (`!`) is encountered, the function should stop processing and return the string up to this point. Any characters after the exclamation sign should be discarded.

Once you are done, try out your new `remove_vowels()` function for `input_string = 'No vowels here! Python is fun!'` in the following cell. Assign the output to `q7`. Feel free to try it for other inputs.

```
    Examples:
    
    >>> remove_vowels('Welcome to E7!')
    'Wlcm t 7!'
    
    >>> remove_vowels('Hello World!')
    'Hll Wrld!'
    
    >>> remove_vowels('No vowels here! Python is fun!')
    'N vwls hr!'
    
    >>> remove_vowels('The quick brown fox jumps over the lazy dog.')
    'Th qck brwn fx jmps vr th lzy dg.'
```
    

```python
# ANSWER CELL
...
```

```python
# TEST YOUR FUNCTION HERE
q7 = ...
```

```python
grader.check("q7")
```

## (Optional) Question 8: Summation Target

**Note: This problem is NOT graded and is provided solely for additional practice. Skipping or completing this problem will not affect your grade.**

Write a function that takes a list of numbers (`myList`) and a single desired value to serve as a target summation (`sumTarget`). Your function needs to find all unique three elements $a, b, c$ in `myList` such that $a + b + c$ = `sumTarget`. 

Your function `Summation(myList, sumTarget)` should return `results`, where `results` is a NumPy array with shape (M, 3), where each row includes unique three elements $a, b, c$ in `myList` that add up to `sumTarget`. 

Constraints:
* In `results`, the three elements $a, b, c$ within each row should be sorted in increasing order. 
* The order of the rows in `results` does not matter.
* No duplicate rows should be present in `results` (each row should be unique).
* If no three elements add up to `sumTarget`, the function should return an empty array: `array([])`.
* Within each row, the same element cannot be selected more than once, unless it appears in `myList` more than once (i.e., selecting elements is done without replacement).

Test your `mySummation()` function for `array = [0, 1, 1, 2, 3, 4]` and `sumTarget = 5`. Assign the output to `q8`. Feel free to try it for other inputs.

*Hint:* You may find [`numpy.vstack()`](https://numpy.org/doc/stable/reference/generated/numpy.vstack.html) and [`numpy.empty()`](https://numpy.org/doc/stable/reference/generated/numpy.empty.html) useful for creating your `results` array.

```
Examples:

>>> mySummation([1, 1, 5, 6, 7, 8, 10], 10)
array([[1., 1., 8.]])

>>> mySummation([0, 1, 1, 2, 3, 4], 5)
array([[0., 1., 4.],
       [0., 2., 3.],
       [1., 1., 3.]])
       
>>> mySummation([0, 1, 1, 2, 3, 4], 10)
array([], dtype=float64)

>>> mySummation([0.1, 0.5, -1, 2, 4, 3.3], -0.4)
array([[-1. ,  0.1,  0.5]])
```

```python
# ANSWER CELL
...
```

```python
# TEST YOUR FUNCTION HERE
q8 = ...

print(q8)
```

```python
grader.check("q8")
```

### You're done with this Lab!

**Important submission information:** After completing the assignment, click on the Save icon from the Tool Bar &nbsp;<i class="fa fa-save" style="font-size:16px;"></i>&nbsp;. After saving your notebook, **run the cell with** `grader.check_all()` and confirm that you pass the same tests as in the notebook. Then, **run the final cell** `grader.export()` and click the link to download the zip file. Then, go to Gradescope and submit the zip file to the corresponding assignment. 

**Once you have submitted, stay on the Gradescope page to confirm that you pass the same tests as in the notebook.**

```python
import matplotlib.image as mpimg
img = mpimg.imread('resources/animal.jpg')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congrats on finishing this lab!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 4)
```
>>> 
>>> # Check students tested the function (type is not ellipsis)
>>> assert get_hash(type(q1)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> 
>>> # Check return value type
>>> assert get_hash(type(mySavingsPlan(1000, 0.05, 10000))) == '2203e3c2293eb8c9bda5bcd97acbf884'
```
failure_message: Make sure your function returns an integer.
```
>>> 
>>> # Check simple cases
>>> assert get_hash(mySavingsPlan(1000, 0.05, 10000)) == '642e92efb79421734881b53e1e1b18b6'
>>> assert get_hash(mySavingsPlan(2000, 0.15, 10000)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(mySavingsPlan(500, 0.4, 5000)) == '8f14e45fceea167a5a36dedd4bea2543'
```
```
>>> 
>>> # Check P0 > goal 
>>> assert get_hash(mySavingsPlan(4000, 0.10, 300)) == 'cfcd208495d565ef66e7dff9f98764da'
```
failure_message: What if P0 is greater than the goal? Are additional years needed for saving?
```
>>> 
>>> # Check P0 == goal
>>> assert get_hash(mySavingsPlan(300, 0.10, 300)) == 'cfcd208495d565ef66e7dff9f98764da'
```
failure_message: What if P0 is equal to the goal? Are additional years needed for saving?


Tests for q2.0 (points: 3)
```
>>> 
>>> # Check students tested the function (type is not ellipsis)
>>> assert get_hash(type(q2_0)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> 
>>> # Test for valid input without 0
>>> assert get_hash(Sum(14927421)) == '34173cb38f07f89ddbebc2ac9128303f'
>>> assert get_hash(Sum(1371)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(Sum(629)) == '70efdf2ec9b086079795c442636b55fb'
```
```
>>> 
>>> # Test for valid input with 0
>>> assert get_hash(Sum(104927421)) == '34173cb38f07f89ddbebc2ac9128303f'
>>> assert get_hash(Sum(13710)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(Sum(60020090)) == '70efdf2ec9b086079795c442636b55fb'
```
failure_message: What if n has 0?
```
>>> 
>>> error_message = ['INVALID', 'INPUT']
>>> # Test for non-integer (use .upper() to avoid failed test due to upper/lower case)
>>> assert get_hash(all([word in Sum(0.1).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in Sum(10.1).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in Sum(10.00000001).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in Sum(10.9).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in Sum(10.0).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: What if n is non-integer?
```
>>> 
>>> error_message = ['INVALID', 'INPUT']
>>> # Test for negative values (use .upper() to avoid failed test due to upper/lower case)
>>> assert get_hash(all([word in Sum(-1).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in Sum(-99).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in Sum(-0.01).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: What if n is negative?


Tests for q2.1 (points: 2)
```
>>> 
>>> # Check students tested the function (type is not ellipsis)
>>> assert get_hash(type(q2_1)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> 
>>> assert callable(lambdaSum) and lambdaSum.__name__ == "<lambda>"
>>> 
>>> # Check if Sum() is being called inside lambdaSum()
>>> Sum_temp = Sum
>>> try:
...     Sum = sum # modify Sum
...     assert lambdaSum(10) == 1
... except TypeError:
...     Sum = Sum_temp
...     assert 0 == 1
>>> Sum = Sum_temp
```
failure_message: Make sure it is a lambda function!
```
>>> 
>>> # Test for valid input without 0
>>> assert get_hash(lambdaSum(14927421)) == '34173cb38f07f89ddbebc2ac9128303f'
>>> assert get_hash(lambdaSum(1371)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(lambdaSum(629)) == '70efdf2ec9b086079795c442636b55fb'
>>> assert callable(Sum) and lambdaSum.__name__ == "<lambda>"
```
```
>>> 
>>> # Test for valid input with 0
>>> assert get_hash(lambdaSum(104927421)) == '34173cb38f07f89ddbebc2ac9128303f'
>>> assert get_hash(lambdaSum(13710)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(lambdaSum(60020090)) == '70efdf2ec9b086079795c442636b55fb'
>>> assert callable(Sum) and lambdaSum.__name__ == "<lambda>"
```
failure_message: What if n has 0?
```
>>> 
>>> error_message = ['INVALID', 'INPUT']
>>> # Test for non-integer (use .upper() to avoid failed test due to upper/lower case)
>>> assert get_hash(all([word in lambdaSum(0.1).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in lambdaSum(10.1).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in lambdaSum(10.00000001).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in lambdaSum(10.9).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in lambdaSum(10.0).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert callable(Sum) and lambdaSum.__name__ == "<lambda>"
```
failure_message: What if n is non-integer?
```
>>> 
>>> error_message = ['INVALID', 'INPUT']
>>> # Test for negative values (use .upper() to avoid failed test due to upper/lower case)
>>> assert get_hash(all([word in lambdaSum(-1).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in lambdaSum(-99).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in lambdaSum(-0.01).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert callable(Sum) and lambdaSum.__name__ == "<lambda>"
```
failure_message: What if n is negative?


Tests for q2.2 (points: 3)
```
>>> 
>>> # Check students tested the function (type is not ellipsis)
>>> assert get_hash(type(q2_2)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> 
>>> # Test for valid input without 0
>>> assert get_hash(SumSize(14927421)) == '34173cb38f07f89ddbebc2ac9128303f'
>>> assert get_hash(SumSize(1371)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(SumSize(629)) == '70efdf2ec9b086079795c442636b55fb'
>>> assert get_hash(SumSize(104927421)) == '34173cb38f07f89ddbebc2ac9128303f'
>>> assert get_hash(SumSize(13710)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(SumSize(60020090)) == '70efdf2ec9b086079795c442636b55fb'
```
failure_message: Make sure the default value of size is set to 1
```
>>> 
>>> assert get_hash(SumSize(14927421,2)) == '757b505cfd34c64c85ca5b5690ee5293'
>>> assert get_hash(SumSize(1371,2)) == '68d30a9594728bc39aa24be94b319d21'
>>> assert get_hash(SumSize(629,3)) == '051e4e127b92f5d98d3c79b195f2b291'
>>> assert get_hash(SumSize(10492742,4)) == '690bb330e5e7e3e07867fafc4d32ec82'
>>> assert get_hash(SumSize(600200900,3)) == '01e00f2f4bfcbb7505cb641066f2859b'
```
failure_message: Make sure the function works for any value of size
```
>>> 
>>> error_message = ['NUMBER', 'DIGITS', 'N', 'SHOULD', 'MULTIPLE', 'SIZE']
>>> assert all([word in SumSize(1,3).upper() for word in error_message])
>>> assert all([word in SumSize(1,2).upper() for word in error_message])
>>> assert all([word in SumSize(12345,3).upper() for word in error_message])
>>> assert all([word in SumSize(15487411,6).upper() for word in error_message])
```
failure_message: What if the number of digits in n is not a multiple of size?


Tests for q3 (points: 3)
```
>>> 
>>> # Check students tested the function (type is not ellipsis)
>>> assert get_hash(type(q3)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> 
>>> # Check output types and length
>>> assert get_hash(type(collatz(4))) == '22b3afce62075c7012f8e5041adfee16'
```
failure_message: Check output format. It should be a tuple.
```
>>> 
>>> # Check output types and length
>>> assert get_hash(type(collatz(4)[0])) == '425ed5b515f0a67c6316f0ec499f86bf'
```
failure_message: Make sure the first returned item is a list.
```
>>> 
>>> # Check output types and length
>>> assert get_hash(len(collatz(4))) == 'c81e728d9d4c2f636f067f89cc14862c'
```
failure_message: Make sure you return both the sequence and the steps.
```
>>> 
>>> # Check n == 1
>>> assert get_hash(int(sum(collatz(1)[0]))) == 'c4ca4238a0b923820dcc509a6f75849b'
>>> assert get_hash(collatz(1)[1]) == 'c4ca4238a0b923820dcc509a6f75849b'
```
failure_message: What if n is 1?
```
>>> 
>>> # Check general cases
>>> assert get_hash(int(sum(collatz(17)[0]))) == 'ca46c1b9512a7a8315fa3c5a946e8265'
>>> assert get_hash(collatz(17)[1]) == 'c51ce410c124a10e0db5e4b97fc2af39'
>>> assert get_hash(int(sum(collatz(119)[0]))) == 'e951ccd95572a67138f4572c1c7d7ee8'
>>> assert get_hash(collatz(119)[1]) == 'e369853df766fa44e1ed0ff613f563bd'
>>> assert get_hash(int(sum(collatz(407)[0]))) == '2af1e9eaa807096a11b32ed26ecc2cbb'
>>> assert get_hash(int(sum(collatz(791)[0]))) == '3e5c39e6359d13d57826c85e119ed584'
```
failure_message: Check the sequence and step count.
```
>>> 
>>> error_message = ['INVALID', 'INPUT']
>>> # Check error message
>>> assert get_hash(all([word in collatz(-1).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: What if n is negative? Check the formatting of the error message.
```
>>> 
>>> error_message = ['INVALID', 'INPUT']
>>> # Check error message
>>> assert get_hash(all([word in collatz(0).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: What if n is zero? Check the formatting of the error message.
```
>>> 
>>> error_message = ['INVALID', 'INPUT']
>>> # Check error message
>>> assert get_hash(all([word in collatz(8.1).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in collatz(0.5).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(all([word in collatz(1.0).upper() for word in error_message])) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: What if n is non-integer? Check the formatting of the error message.


Tests for q4 (points: 4)
```
>>> 
>>> # Check students tested the function (type is not ellipsis)
>>> assert get_hash(type(q4)) != '14e736438b115821cbb9b7ac0ba79034', "Make sure to test your function!"
```
failure_message: Make sure to test your function!
```
>>> 
>>> # Check output structure
>>> assert get_hash(type(Rainfall([2.3,4.1,-5.4,1.0,7.1]))) == '22b3afce62075c7012f8e5041adfee16'
```
failure_message: Check output format. It should be a tuple.
```
>>> 
>>> # Check  length
>>> assert get_hash(len(Rainfall([2.3,4.1,-5.4,1.0,7.1]))) == 'c81e728d9d4c2f636f067f89cc14862c'
```
failure_message: Make sure you return both the total and average rainfall.
```
>>> 
>>> # Check simple case
>>> assert get_hash(round(Rainfall([2.3,4.1,5.4,1.0,7.1,2,8.0,4.9,6])[0], 3)) == 'a488d325d66dc4190e9637e240990868'
>>> assert get_hash(round(Rainfall([20, 3, 0, 23, 8, 7, 24, 4, 28, 17])[0], 3)) == '02522a2b2726fb0a03bb19f2d8d9524d'
```
failure_message: Check rainTot
```
>>> 
>>> # Check simple case
>>> assert get_hash(round(Rainfall([2.3,4.1,5.4,1.0,7.1,2,8.0,4.9,6])[1], 3)) == '7cc0ede3573439b052e865e1b4213db4'
>>> assert get_hash(round(Rainfall([20, 3, 0, 23, 8, 7, 24, 4, 28, 17])[1], 3)) == 'dfd78eb85404bb841184ad07c266066f'
```
failure_message: Check rainAvg
```
>>> 
>>> # Check simple case
>>> assert get_hash(round(Rainfall([2.3,4.1,5.4,1.0,7.1,2,8.0,0,4.9,6])[0], 3)) == 'a488d325d66dc4190e9637e240990868'
>>> assert get_hash(round(Rainfall([20, 3, 0, 23, 8, 7, 24, 4, 0, 28, 17])[0], 3)) == '02522a2b2726fb0a03bb19f2d8d9524d'
```
failure_message: Check rainTot when there is a measurement equal to 0
```
>>> 
>>> # Check simple case
>>> assert get_hash(round(Rainfall([2.3,4.1,5.4,1.0,7.1,2,8.0,0,4.9,6])[1], 3)) == '6187c0685f2fe6d1b2595c847b124663'
>>> assert get_hash(round(Rainfall([20, 3, 0, 23, 8, 7, 24, 4, 0, 28, 17])[1], 3)) == 'bfd1ace15fae5dc2325cbda285874522'
```
failure_message: Check rainTot when there is a measurement equal to 0
```
>>> 
>>> # Check negative numbers
>>> assert get_hash(round(Rainfall([-2, -0.5, 1, 0, 4])[0],3)) == 'e4da3b7fbbce2345d7772b0674a318d5'
>>> assert get_hash(round(Rainfall([-2, -0.5, 1, 0, 4])[1],3)) == '2280082248f35fad32c7ead135ab3f99'
```
failure_message: What if the list contains negative numbers?
```
>>> 
>>> # Check negative numbers
>>> assert get_hash(round(Rainfall([-13,-19,-3,-15,7,-20,-18,13,9,17,11,0,2,-7,6,15,14])[0], 3)) == 'f4b9ec30ad9f68f89b29639786cb62ef'
>>> assert get_hash(round(Rainfall([-13,-19,-3,-15,7,-20,-18,13,9,17,11,0,2,-7,6,15,14])[1], 3)) == 'b68de8f1f938e0330acdf64c63b08b9a'
```
failure_message: What if the list contains negative numbers?
```
>>> 
>>> # Check cases with early stop
>>> assert get_hash(round(Rainfall([2,1.3,-99,33,0.99,-999])[0], 3)) == '652126365f55aa7ccf24e90ed647c3a3'
>>> assert get_hash(round(Rainfall([2,1.3,-99,-999,33,0.99,1])[0], 3)) == 'bca66716b9ca2d13a4c28be69b276777'
>>> assert get_hash(round(Rainfall([0, -999,1.3,-99,33,0.99,-999])[0], 3)) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(round(Rainfall([2,1.3,-99,33,0.99,-999])[1], 3)) == '5cf35d410e296c02b09670b03e3d38c5'
>>> assert get_hash(round(Rainfall([2,1.3,-99,-999,33,0.99,1])[1], 3)) == '448258b68b4ebf3a6590c75fca600c29'
>>> assert get_hash(round(Rainfall([0, -999,1.3,-99,33,0.99,-999])[1], 3)) == '30565a8911a6bb487e3745c0ea3c8224'
```
failure_message: What if the list contains -999?


Tests for q5 (points: 3)
```
>>> 
>>> # Check students tested the function (type is not ellipsis)
>>> assert get_hash(type(q5)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> 
>>> import inspect
>>> assert any([True if GenRec.__name__ in line and '(' in line else False for line in inspect.getsourcelines(GenRec)[0][1:]])
>>> assert all([False if 'for' in line else True for line in inspect.getsourcelines(GenRec)[0][1:]])
>>> assert all([False if 'while' in line else True for line in inspect.getsourcelines(GenRec)[0][1:]])
```
failure_message: Your function should use recursion! for and while keywords are not allowed in the function (not even in comments)
```
>>> 
>>> import inspect
>>> assert any([True if GenRec.__name__ in line and '(' in line else False for line in inspect.getsourcelines(GenRec)[0][1:]])
>>> assert all([False if 'for' in line else True for line in inspect.getsourcelines(GenRec)[0][1:]])
>>> assert all([False if 'while' in line else True for line in inspect.getsourcelines(GenRec)[0][1:]])
>>> 
>>> # Test base-cases
>>> assert get_hash(GenRec(2, 4, 0, 2, 0)) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(GenRec(2, 4, 1, 2, 0)) == 'c4ca4238a0b923820dcc509a6f75849b'
>>> assert get_hash(GenRec(4, 2, 2, 1, 0)) == 'c81e728d9d4c2f636f067f89cc14862c'
```
failure_message: Check the base case for n=0
```
>>> 
>>> import inspect
>>> assert any([True if GenRec.__name__ in line and '(' in line else False for line in inspect.getsourcelines(GenRec)[0][1:]])
>>> assert all([False if 'for' in line else True for line in inspect.getsourcelines(GenRec)[0][1:]])
>>> assert all([False if 'while' in line else True for line in inspect.getsourcelines(GenRec)[0][1:]])
>>> 
>>> # Test base-cases
>>> assert get_hash(GenRec(4, 3, 0, 2, 1)) == 'c81e728d9d4c2f636f067f89cc14862c'
>>> assert get_hash(GenRec(5, 1, 1, 3, 1)) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
>>> assert get_hash(GenRec(0, 1, 2, 1, 1)) == 'c4ca4238a0b923820dcc509a6f75849b'
```
failure_message: Check the base case for n=1
```
>>> 
>>> import inspect
>>> assert any([True if GenRec.__name__ in line and '(' in line else False for line in inspect.getsourcelines(GenRec)[0][1:]])
>>> assert all([False if 'for' in line else True for line in inspect.getsourcelines(GenRec)[0][1:]])
>>> assert all([False if 'while' in line else True for line in inspect.getsourcelines(GenRec)[0][1:]])
>>> 
>>> # Test for valid input n > 2
>>> assert get_hash(GenRec(4, 3, 0, 2, 5)) == 'f0adc8838f4bdedde4ec2cfad0515589'
>>> assert get_hash(GenRec(4, 3, 1, 2, 5)) == '818f4654ed39a1c147d1e51a00ffb4cb'
>>> assert get_hash(GenRec(1, 2, 0, 0, 5)) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(GenRec(0, 0, 1, 5, 10)) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(GenRec(1, 1, 1, 5, 10)) == 'eddea82ad2755b24c4e168c5fc2ebd40'
>>> assert get_hash(GenRec(1, 2, 1, 5, 25)) == '6938db3189d0daee0bd06d4d54429da4'
```
failure_message: Check the general case


Tests for q6.0 (points: 0.5)
```
>>> 
>>> assert get_hash(type(round(q6a))) == '2203e3c2293eb8c9bda5bcd97acbf884'
```
failure_message: Your answer should be numeric.
```
>>> assert get_hash(round(q6a)) == 'cfcd208495d565ef66e7dff9f98764da'
```


Tests for q6.1 (points: 0.5)
```
>>> 
>>> assert len(q6b) == 1
>>> assert isinstance(q6b, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format. Note that the actual test is hidden.


Tests for q6.2 (points: 0.5)
```
>>> 
>>> assert len(q6c) == 1
>>> assert isinstance(q6c, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format. Note that the actual test is hidden.


Tests for q6.3 (points: 0.5)
```
>>> 
>>> assert len(q6d) == 1
>>> assert isinstance(q6d, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format. Note that the actual test is hidden.


Tests for q6.4 (points: 1)
```
>>> 
>>> # Test base-cases
>>> assert get_hash(kbonacci(0,2)) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(kbonacci(0,3)) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(kbonacci(1,4)) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(kbonacci(2,5)) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(kbonacci(3,5)) == 'cfcd208495d565ef66e7dff9f98764da'
```
failure_message: Check the base case
```
>>> 
>>> # Test base-cases
>>> assert get_hash(kbonacci(1,2)) == 'c4ca4238a0b923820dcc509a6f75849b'
>>> assert get_hash(kbonacci(2,3)) == 'c4ca4238a0b923820dcc509a6f75849b'
>>> assert get_hash(kbonacci(3,4)) == 'c4ca4238a0b923820dcc509a6f75849b'
>>> assert get_hash(kbonacci(4,5)) == 'c4ca4238a0b923820dcc509a6f75849b'
```
failure_message: Check the base case
```
>>> 
>>> assert get_hash(kbonacci(5,2)) == 'e4da3b7fbbce2345d7772b0674a318d5'
>>> assert get_hash(kbonacci(10,2)) == 'b53b3a3d6ab90ce0268229151c9bde11'
>>> assert get_hash(kbonacci(10,3)) == '43ec517d68b6edd3015b3edc9a11367b'
>>> assert get_hash(kbonacci(10,4)) == '9f61408e3afb633e50cdf1b20de6f466'
>>> assert get_hash(kbonacci(5,3)) == 'a87ff679a2f3e71d9181a67b7542122c'
>>> assert get_hash(kbonacci(4,4)) == 'c4ca4238a0b923820dcc509a6f75849b'
>>> assert get_hash(kbonacci(15,5)) == '2a9d121cd9c3a1832bb6d2cc6bd7a8a7'
>>> assert get_hash(kbonacci(14,6)) == '621bf66ddb7c962aa0d22ac97d69b793'
```
failure_message: Check the general case


Tests for q7 (points: 0)
```
>>> 
>>> # Check students tested the function (type is not ellipsis)
>>> assert get_hash(type(q7)) != '14e736438b115821cbb9b7ac0ba79034', "Make sure to test your function!"
```
failure_message: Make sure to test your function!
```
>>> # Check examples
>>> assert get_hash(remove_vowels('Welcome to E7!')) == 'c796d49406d6eaceb4b8f0a60046d990'
>>> assert get_hash(remove_vowels('Hello World!')) == '21342b09fe42a27da8720d77704ccd73'
>>> assert get_hash(remove_vowels('No vowels here! Python is fun!')) == '68a087a3a7b8ff6d94aa9226b6d0f5a3'
>>> assert get_hash(remove_vowels('The quick brown fox jumps over the lazy dog.')) == 'f80f4502beacd9a425caef43761e9047'
```
```
>>> # Check new strings
>>> assert get_hash(remove_vowels('AEIOUaeiou')) == 'd41d8cd98f00b204e9800998ecf8427e'
>>> assert get_hash(remove_vowels('AEIOUaeiou! The end.')) == '9033e0e305f247c0c3c80d0c7848c8b3'
>>> assert get_hash(remove_vowels('Programming is fun!!!')) == '6f721ade8c0b0f296b46a4990f94a1f6'
>>> assert get_hash(remove_vowels('One two three four five')) == 'fbd35e565eb9e36330960a16c429b7d2'
```


Tests for q8 (points: 0)
```
>>> 
>>> # Check students tested the function (type is not ellipsis)
>>> assert get_hash(type(q8)) != '14e736438b115821cbb9b7ac0ba79034', "Make sure to test your function!"
```
failure_message: Make sure to test your function!
```
>>> 
>>> # Check data structure
>>> assert get_hash(type(mySummation([1, 1, 5, 6, 7, 8, 10],10))) == '30b39aca73d0815f5ca1a06af1412947'
>>> assert get_hash(mySummation([1, 1, 5, 6, 7, 8, 10],10).shape) == '0f6c6ad8083308c6c59c917667b3fa2f'
```
failure_message: Check output type.
```
>>> 
>>> # Check data structure
>>> assert get_hash(type(mySummation([1, 1, 5, 6, 7, 8, 10],10))) == '30b39aca73d0815f5ca1a06af1412947'
>>> assert get_hash(mySummation([1, 1, 5, 6, 7, 8, 10],10).shape) == '0f6c6ad8083308c6c59c917667b3fa2f'
```
failure_message: Check output shape.
```
>>> 
>>> # Check empty cases
>>> assert get_hash(mySummation([0, 1, 1, 2, 3, 4],10)) == 'd751713988987e9331980363e24189ce'
>>> assert get_hash(mySummation([10,20,30,40,50],200)) == 'd751713988987e9331980363e24189ce'
```
failure_message: What if no three elements add up to sumTarget?
```
>>> 
>>> assert get_hash(mySummation([10,0,1,2],21)) == 'd751713988987e9331980363e24189ce'
```
failure_message: Make sure to not select an element more than once.
```
>>> 
>>> # Check simple case
>>> assert get_hash(mySummation([2,3,4,5,6],12).shape) == '590392f0cb88f1eb89930f4de5e9eaaa'
>>> assert get_hash(int(sum(mySummation([2,3,4,5,6],12)[:,2]))) == '6512bd43d9caa6e02c990b0a82652dca'
```
failure_message: Make sure your rows are sorted.
```
>>> 
>>> # Check case with multiple occurence
>>> assert get_hash(int(sum(mySummation([0, 0, 2, 4, 6, 8, 14],14)[:,2]))) == '34173cb38f07f89ddbebc2ac9128303f'
>>> assert get_hash(int(sum(mySummation([1,5,6,5,1],12)[:,2])))
```
failure_message: Make sure all rows are unique.
```
>>> 
>>> # Check random case 
>>> arr = [3, -4, 2, -2, -7, -10, 4, 5, 8, 6, 7, -6]
>>> assert get_hash(int(sum(mySummation(arr,10)[:,2]))) == '33e75ff09dd601bbe69f351039152189'
>>> arr = [9, 4, 5, -4, 2, -6, 7, -9, 8, -7, 1, -2]
>>> assert get_hash(int(sum(mySummation(arr,10)[:,2]))) == 'a684eceee76fc522773286a895bc8436'
```
failure_message: Make sure all rows are unique.
